# Introduction to Data Science

```python
from itertools import combinations
from sklearn.datasets import load_iris
import matplotlib.pyplot as plt
import pandas as pd
import seaborn as sns
from sklearn.decomposition import PCA
import viz_helper as viz
import matplotlib.cm as cm
import numpy as np
%matplotlib inline
```

```python
IRIS = load_iris()
```

```python
names = [(name
  .replace(" ", "")
  .replace("(cm)", "")) for name in IRIS.feature_names]
```

```python
iris_df = pd.DataFrame(IRIS.data)
iris_df.columns = names
```

```python
iris_df.head()
```

```python
fig = plt.figure(figsize=(20,3))
for i, col in enumerate(iris_df.columns):
    fig.add_subplot(1,4,1+i)
    sns.distplot(iris_df[col])
    plt.xlim(-3,3)
```

```python
from sklearn.preprocessing import StandardScaler

scaler = StandardScaler()
iris_scaled = scaler.fit_transform(iris_df)
iris_scaled_df = pd.DataFrame(iris_scaled, columns=names)
```

```python
iris_scaled_df.head()
```

```python
fig = plt.figure(figsize=(20,6))
for i, col in enumerate(iris_df.columns):
    fig.add_subplot(2,4,1+i)
    sns.distplot(iris_df[col])
    plt.xlim(-3,3)
    fig.add_subplot(2,4,5+i)
    sns.distplot(iris_scaled_df[col])
    plt.xlim(-3,3)
```

```python
fig = plt.figure(figsize=(20,8))
for i, pair in enumerate(combinations(names, 2)):
    name_1 = pair[0]
    name_2 = pair[1]
    feat_1 = iris_scaled_df[name_1]
    feat_2 = iris_scaled_df[name_2]
    fig.add_subplot(2,3,1+i)
    plt.scatter(feat_1, feat_2, c=IRIS.target, cmap=cm.viridis)
    plt.xlabel(name_1)
    plt.ylabel(name_2)
```

```python
pca = PCA()
iris_pca = pca.fit_transform(iris_scaled_df)
iris_pca_df = pd.DataFrame(iris_pca, columns=['PC 1', 'PC 2', 'PC 3', 'PC 4'])
```

```python
iris_pca_df.head()
```

```python
viz.pca_results(iris_scaled_df, pca);
```

```python
plt.scatter(iris_pca_df['PC 1'], iris_pca_df['PC 2'], c=IRIS.target, cmap=cm.viridis)
plt.xlabel('PC 1')
plt.ylabel('PC 2')
```

<!-- grade_id: cell-12e29d3e63967406, points: 10, grade: true, solution: false -->
```python
assert len(IRIS) == 5
```

```python

```